# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
from numpy import nan as NA
```

### 4.2 Escreva uma função para calcular os percentis [0.25, 0.50, 0.75, 0.9] do dataset do índice de gini dos municípios brasileiros para todos os anos presents no dataset.

#### Apenas por curiosidade, plote o gráfico destes percentis num gráfico de linhas


<!-- grade_id: cell-553qe5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Descubra qual a string representa valores NA neste dataset.
# Dica 2: Use o argumento na_values para especificar a string que representa valores NA.
# Dica 3: Use o argumento decimal=',' para definir a vírgula como separador de decimal.
# Dica 4: Use a função apply e uma função lambda para desmembrar a coluna município no
#         respectivo nome e código do município.
# Dica 5: Atente para o separador de colunas do dataset.
# Dica 6: O encoding do dataset do gini é 'utf8'.
# Dica 7: Atente para quantas linhas você deve ignorar antes no inicio e no final do arquivo.
# Dica 8: Use a função quantile
def plot_gini(path_gini):
    """retorne o valor de: n+nn+nnn
    """
    # YOUR CODE HERE
    df = pd.read_csv(path_gini, sep=";", skiprows=2, header=0, na_values="...", decimal=",")
    
    return df[["1991", "2000", "2010"]].quantile(q=[0.25, 0.5, 0.75, 0.9])

```

```python
## Faça seus testes aqui
# Considere o seguinte dataframe do PIB dos municípios
path_gini = '../../curso_ciencia_dados2020/data/originais/gini/ginibr.csv'

plot_gini(path_gini).apply(lambda x: round(x, 2)).to_dict()

```

Output:
```
{'1991': {0.25: 0.49, 0.5: 0.53, 0.75: 0.58, 0.9: 0.62},
 '2000': {0.25: 0.51, 0.5: 0.55, 0.75: 0.6, 0.9: 0.64},
 '2010': {0.25: 0.46, 0.5: 0.5, 0.75: 0.55, 0.9: 0.58}}
```

<!-- grade_id: cell-0e2bh6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
test_data = {'1991': {0.25: 0.49, 0.5: 0.53, 0.75: 0.58, 0.9: 0.62},
 '2000': {0.25: 0.51, 0.5: 0.55, 0.75: 0.6, 0.9: 0.64},
 '2010': {0.25: 0.46, 0.5: 0.5, 0.75: 0.55, 0.9: 0.58}}
path_gini = '../../curso_ciencia_dados2020/data/originais/gini/ginibr.csv'

df = plot_gini(path_gini)
assert df.apply(lambda x: round(x, 2)).to_dict() == test_data
```

```python

```