# Aula 3 - Exercícios - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np
import sqlalchemy as sqla

```

### 3.5 Escreva uma função para calcular o PIB percapita de cada município do Brasil, considerando o dataframe do exercício anterior.
#### Nomeie a coluna do PIB Percapita com 'pib_percapita'
#### Renomeie a coluna municipio_x para municipio

<!-- grade_id: cell-553e524r8eb14bc30, points: 0, grade: false, solution: true -->
```python

def pib_percapita(pib_pop):
    """retorne retorne um dataframe contendo o codigo ibge de 7 caracteres,
       o nome do municipio e seu PIB percapita.
    """
    # YOUR CODE HERE
    
    pib_pop["pib_percapita"] = pib_pop['2017'] / pib_pop['pop2017']
    return pib_pop
```

```python
# Importar e carregar as funcoes do exercicio 4
%run ../Aula3_Ex04/Aula3_Ex04.ipynb
```

```python
## Faça seus testes aqui

df_pip_pop = merge_pop_pib(path_pop, path_pib)
pib_percapita(df_pip_pop)

```

Output:
```
nivel cod_ibge7                 municipio_x      2007       2009  \
0       MU   1100015  Alta Floresta D'Oeste (RO)    191364     256986   
1       MU   1100023              Ariquemes (RO)    905203    1133095   
2       MU   1100031                 Cabixi (RO)     49166      69776   
3       MU   1100049                 Cacoal (RO)    814890     985479   
4       MU   1100056             Cerejeiras (RO)    143270     190902   
...    ...       ...                         ...       ...        ...   
5565    MU   5222005             Vianópolis (GO)    122694     164336   
5566    MU   5222054         Vicentinópolis (GO)     87881     133430   
5567    MU   5222203               Vila Boa (GO)     31897      73367   
5568    MU   5222302          Vila Propício (GO)     76868     107125   
5569    MU   5300108               Brasília (DF)  93404001  124323753   

           2011       2013       2015       2017  uf  cod_uf  cod_munic  \
0        280510     341325     421300     498864  RO 
```

<!-- grade_id: cell-0e2hb6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
# Importar e carregar as funcoes do exercicio 4
%run ../Aula3_Ex04/Aula3_Ex04.ipynb
df_pib_pop = merge_pop_pib(path_pop, path_pib)
pib_percapita_2017 = pib_percapita(df_pib_pop)
assert round(pib_percapita(df_pib_pop).loc[0,'pib_percapita'],2) == 19.61
```

```python

```