# Aula 6 - Exercícios - Storytelling e Correlação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt # Biblioteca para plotar gráficos
import seaborn as sns # Biblioteca para plotar gráficos baseada no matplotlib
import re
```

### 6.1 Escreva uma função para encontrar numa string todos os grupos de letras em que uma letra a é seguida de 3 letras b. 


<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
def a3b(s):
    """retorne o valor de: n+nn+nnn
    """
    # YOUR CODE HERE
    return re.findall('[a-z]b{3}',s)

```

```python
## Faça seus testes aqui
a3b('abbb asde ababbaaabbbabbb') 

```

Output:
```
['abbb', 'abbb', 'abbb']
```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert a3b('abbb asde ababbaaabbbabbb') == ['abbb', 'abbb', 'abbb']
```

```python

```