# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
```

### 4.1 Escreva uma função para ler os dados de PIB e população dos datasets disponibilizados e retornar um dataframe com a agregação por Estados e as métricas soma, média e mediana do PIB de 2017 e da população de todos os Estados brasileiros.
#### Você pode se basear na sua resolução dos exercícios da Aula 3.


<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica: Baseie seu código na sua solução do exercício da aula 3 que faz o merge dos
#       datasets do PIB e da população.

def aggregate_pib_pop(path_pop, path_pib):
    """retorne um dataframe com a agregação por Estados e as métricas soma, média, mediana
       e o desvio padrão do PIB de 2017 e da população de todos os Estados brasileiros.
    """
    # YOUR CODE HERE
    df_pib = pd.read_excel(path_pib, sheet_name='Tabela', skiprows=3, skipfooter=1, header=0) 
    df_pop = pd.read_excel(path_pop, sheet_name="Municipios", header=0)
    df_pib.rename(columns={"Unnamed: 0":"nivel", "Unnamed: 1": "cod_ibge7", "Unnamed: 2":"municipio"}, inplace=True)

    df_pop["cod_ibge7"] = df_pop["cod_uf"].astype(dtype='string') + df_pop['cod_munic'].astype(dtype="string").str.zfill(5)
    df_pib["cod_ibge7"] = df_pib["cod_ibge7"].astype(dtype="string")

    df_pip_pop = pd.merge(df_pop, df_pib, on="cod_ibge7")
    
    grouped = df_pip_pop[["2017","pop2017"]].groupby(df_pip_pop["uf"])
    
    
    return grouped.agg(['mean','std', 'median', 'sum'])


```

```python
## Faça seus testes aqui
path_pib = '../../curso_ciencia_dados2020/data/originais/pib/pib_municipios.xlsx'
path_pop='../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'

aggregate_pib_pop(path_pop, path_pib).loc['AC'].apply(lambda x: round(x,1)).to_dict()

```

Output:
```
{('2017', 'mean'): 648684.6,
 ('2017', 'std'): 1695166.9,
 ('2017', 'median'): 229775.0,
 ('2017', 'sum'): 14271061.0,
 ('pop2017', 'mean'): 37710.0,
 ('pop2017', 'std'): 79036.0,
 ('pop2017', 'median'): 17903.5,
 ('pop2017', 'sum'): 829619.0}
```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_pib = '../../curso_ciencia_dados2020/data/originais/pib/pib_municipios.xlsx'
path_pop='../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
test_data = {('pop2017', 'sum'): 829619.0,
 ('pop2017', 'mean'): 37710.0,
 ('pop2017', 'median'): 17903.5,
 ('pop2017', 'std'): 79036.0,
 ('2017', 'sum'): 14271061.0,
 ('2017', 'mean'): 648684.6,
 ('2017', 'median'): 229775.0,
 ('2017', 'std'): 1695166.9}
results_ac = aggregate_pib_pop(path_pop, path_pib).loc['AC'].apply(lambda x: round(x,1))

assert results_ac.to_dict() == test_data
```

```python

```