# Aula 2 - Exercícios Extra - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas e numpy
import pandas as pd
import numpy as np
```

### 2.8 Escreva uma função para ordenar uma lista de dicionários usando uma função Lambda. 
### A função criada deve receber como argumento a chave do dicionário que será usada 
### na ordenação da lista.


<!-- grade_id: cell-c705b9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica: use a função sorted
def sort_dict_list(l, col):
    """retorne os valores do complemento de s1 em relação a s2
    """
    # YOUR CODE HERE
    
    l.sort(key=lambda list_element_dict: list_element_dict[col])
    
    return l

```

```python
## Faça seus testes aqui
data = [{'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
        {'name': ['David'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
       {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']}]
sort_dict_list(data, "name")
```

Output:
```
[{'name': ['David'],
  'class': ['V'],
  'subject_integration': ['english', 'math', 'science']},
 {'name': ['Sara'],
  'class': ['V'],
  'subject_integration': ['english', 'math', 'science']},
 {'name': ['Surya'],
  'class': ['V'],
  'subject_integration': ['english', 'math', 'science']}]
```

<!-- grade_id: cell-9ffda88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
data = [{'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
        {'name': ['David'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
       {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']}]

data_sorted = [{'name': ['David'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
        {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
       {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']}]

assert sort_dict_list(data, 'name') == data_sorted
```

```python

```

```python

```