# Aula 1 - Exercícios Extra - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.11 Escreva uma função em Python receba como parâmetro duas listas de numeros inteiros e retorne os elementos que são comuns entre elas (sem repetição). Use a sintaxe inline (list comprehension).
##### Exemplo: a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89] e b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]
##### Resultado esperado: 

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
# Dica 1: aninhe na mesma instrução o loope (for) e a condição (if)
# Dica 2: se ao final sua lista contiver repetições, transforme a lista num conjunto com o
# construtor de conjuntos set(), pois os sets não contêm repetições.
def intersecao(a, b):
    """retorne os valor presentes nas duas listas
    """
    # YOUR CODE HERE
    #return set([n if n in b else None for n in a])
    return list(set([x for x in a if x in b]))
    #raise NotImplementedError()
```

```python
## Faça seus testes aqui
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89] 
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]
#print("sim") if 800 in a else None
#inter = intersecao(a,b) 
#inter.remove(None)
print(intersecao(a,b) )
```

Output:
```
[1, 2, 3, 5, 8, 13]
```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89] 
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]
assert intersecao(a,b) == [1, 2, 3, 5, 8, 13]
```

```python

```