# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.7 Escreva uma função Python para contar a quantidade de itens de um dado valor num array NumPy
#### Exemplo:  para um array([1  4  10  8  4  20  23]) e n=4
#### Resultado esperado: 2
#### Obs.: Lembre-se de usar funções numpy que realizam as operações sem a necessidade de escrever laços (loops)

```python
# import numpy
import numpy as np
```

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
# Dica: A função count_nonzero do pacote numpy conta a quantidade 
#       de elementos diferentes de zero ou a quantidade de elementos True.

def count_value(array, n):
    """retorne a quantidade de vezes que n ocorre 
    """
    # YOUR CODE HERE
    return np.count_nonzero(array == n)
    #raise NotImplementedError()
```

```python
## Faça seus testes aqui
a1 = np.array([1,0,1,9,5,3,1,2,4,5,6,1,2,3,4,8,9,2,1])
count_value(a1, 1)

```

Output:
```
5
```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a1 = np.array([1,0,1,9,5,3,1,2,4,5,6,1,2,3,4,8,9,2,1])
assert count_value(a1, 1) == 5

```

```python

```