# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.3 Escreva uma função contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2


<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
# Dica 2: Uma string é uma lista de caracteres
def contagem(lista):
    """retorne o maior item da lista
    """
    # YOUR CODE HERE
    count = 0
    
    # Iterates over every string in the list
    for string in lista:
        # Tests if the string size is grater than 2 and if the first and the last char are equals
        if len(string) > 2 and string[0] == string[-1]:
            # if the condition is satisfied, the count increases in one
            count = count + 1
    
    # after the loop, returns the total of the count
    return count
    #raise NotImplementedError()
```

```python
## Faça seus testes aqui
# do nothing
'''def contagem(lista):
    count  = 0
    for string in lista:
        if len(string) > 2 and string[0] == string[-1]:
            count = count + 1 

    return count

print(contagem(['aaa', 'bbb', 'ccc']))
assert contagem(['abc', 'xyz', 'aba', '1221']) == 2
assert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0'''

```

Output:
```
"def contagem(lista):\n    count  = 0\n    for string in lista:\n        if len(string) > 2 and string[0] == string[-1]:\n            count = count + 1 \n\n    return count\n\nprint(contagem(['aaa', 'bbb', 'ccc']))\nassert contagem(['abc', 'xyz', 'aba', '1221']) == 2\nassert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0"
```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert contagem(['abc', 'xyz', 'aba', '1221']) == 2
assert contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
assert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
```

```python

```