# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
```

### 4.4 Escreva uma função para agregar o dataset do Titanic por classe (pclass) e sexo (sex) e calcular a soma da coluna survived (sobrevivente) e a média e mediana da coluna fare (preço do ticket).


<!-- grade_id: cell-553em5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica: Faça a agregação usando um dicionário que mapeia a coluna e a(s) respectiva(s) métrica(s) de
#       agregação solicitada(s) no enunciado.
def agg_titanic(df_titanic):
    """retorne aagregação do dataframe df_titanic por pclass e sex e calcular a soma 
       da coluna survived e a média e mediana da coluna fare.
    """
    # YOUR CODE HERE
    grouped = df_titanic.groupby(["pclass", "sex"])
    return grouped.agg({'survived': ['sum'], 'fare':['mean','median']})

```

```python
## Faça seus testes aqui
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'

test_data = {('survived', 'sum'): {(1, 'female'): 139,  (1, 'male'): 61,
  (2, 'female'): 94,  (2, 'male'): 25,  (3, 'female'): 106,  (3, 'male'): 75},
  ('fare', 'mean'): {(1, 'female'): 109.4,  (1, 'male'): 69.9,  (2, 'female'): 23.2,
  (2, 'male'): 19.9,  (3, 'female'): 15.3,  (3, 'male'): 12.4}, 
  ('fare', 'median'): {(1, 'female'): 80.9,  (1, 'male'): 49.5,  
  (2, 'female'): 23.0,  (2, 'male'): 13.0,  (3, 'female'): 10.5,  (3, 'male'): 7.9}}

df_titanic = pd.read_excel(path_titanic)

result_agg = agg_titanic(df_titanic).apply(lambda x: round(x, 1))

result_agg.to_dict() == test_data


```

Output:
```
True
```

<!-- grade_id: cell-0se2b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'
test_data = {('survived', 'sum'): {(1, 'female'): 139,  (1, 'male'): 61,
  (2, 'female'): 94,  (2, 'male'): 25,  (3, 'female'): 106,  (3, 'male'): 75},
  ('fare', 'mean'): {(1, 'female'): 109.4,  (1, 'male'): 69.9,  (2, 'female'): 23.2,
  (2, 'male'): 19.9,  (3, 'female'): 15.3,  (3, 'male'): 12.4}, 
  ('fare', 'median'): {(1, 'female'): 80.9,  (1, 'male'): 49.5,  
  (2, 'female'): 23.0,  (2, 'male'): 13.0,  (3, 'female'): 10.5,  (3, 'male'): 7.9}}

df_titanic = pd.read_excel(path_titanic)
result_agg = agg_titanic(df_titanic).apply(lambda x: round(x, 1))

assert result_agg.to_dict() == test_data
```

```python

```