# Aula 3 - Exercícios - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np
import sqlalchemy as sqla
from numpy import nan as NA

```

### 3.8 Considere algumas linhas de uma base de dados de morbidade apresentadas a seguir. Perceba que há dados de vários meses de 2020 para cada Estado da Federação.
#### Escreva uma função para pivotar esta base de dados apresentando os meses do ano nas colunas, os municípios nas linhas e a quantidade de óbitos como valores das células. 

|       | Município               |   cod_ibge6 | mes   |   ano | mes_ano    | Óbitos   |
|------:|:------------------------|------------:|:------|------:|:-----------|:---------|
|     0 | Alta Floresta D'Oeste   |      110001 | Abr   |  2020 | 2020-04-01 | 8        |
|     1 | Alto Alegre dos Parecis |      110037 | Abr   |  2020 | 2020-04-01 | 1        |
|     2 | Alto Paraíso            |      110040 | Abr   |  2020 | 2020-04-01 | 2        |
| 33347 | Vila Boa                |      522220 | Jan   |  2020 | 2020-01-01 | -        |
| 33348 | Vila Propício           |      522230 | Jan   |  2020 | 2020-01-01 | -        |
| 33349 | Brasília                |      530010 | Jan   |  2020 | 2020-01-01 | 578      |

<!-- grade_id: cell-553e52r48eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Sempre dê preferência para fazer agregação usando códigos como identificador, 
#         use nomes em ultimo caso quando não houver códigos identificadores. 
# Dica 2: Se você usar o atributo 'mes_ano' para pivotar, o pandas já ordena de forma cronológica as colunas.
def pivot_morbidade(file_path):
    """retorne um dataframe pivotado
    """
    # YOUR CODE HERE
    
    #raise NotImplementedError()
    #pivot_morbidade(file_path)
    df = pd.read_csv(file_path, sep=';', header=0, index_col="cod_ibge6")  
    df_clean = df.drop_duplicates(subset=["Município","mes_ano"])
    #df["Óbitos"].replace("-", 0, inplace=True)
    #df["Óbitos"] = df["Óbitos"].astype("int")
    #df.sort_values(by="Município")
    #df2 = df.groupby([df['Município'], df['mes_ano']], as_index=False)['Óbitos'].sum()

    
    return df_clean.pivot(columns="mes_ano", values="Óbitos")
    #df_clean.pivot(, columns="mes_ano", values="Óbitos").loc[110001]
    #df2.pivot(index="Município", columns="mes_ano", values="Óbitos").loc[110001]


```

```python
## Faça seus testes aqui
file_path = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade_simplificado.csv'
pivot_morbidade(file_path).loc[110001].to_dict()

```

Output:
```
{'2020-01-01': '5',
 '2020-02-01': '2',
 '2020-03-01': '3',
 '2020-04-01': '8',
 '2020-05-01': '2',
 '2020-06-01': '2'}
```

<!-- grade_id: cell-0e2b6tcefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
file_path = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade_simplificado.csv'
test_data = {'2020-01-01': '5', '2020-02-01': '2', '2020-03-01': '3', '2020-04-01': '8', '2020-05-01': '2', 
             '2020-06-01': '2'}
assert pivot_morbidade(file_path).loc[110001].to_dict() == test_data

```

```python

```

```python

```