# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
```

### 4.5 Escreva uma função para fazer o merge do dataset de morbidade consolidado com o dataset de população
#### O dataset da população deve ficar a esquerda e o dataset da morbidade a direita no resultado.

<!-- grade_id: cell-553a5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: No dataset de morbidade existem valores NA representados por caracteres específicos em várias colunas 
#         do dataframe. Descubra quais são esses caracteres e especifiquei no argumento
#         na_values. Passe nesse argumento uma lista como os caracteres que devem ser 
#         interpretados como NA.
# Dica 2: No dataset de morbidade, especifique o separador, o caracter de números decimais e o tipo do dado do 
#         cod_ibge6 como str. Ou seja, dtype={'cod_ibge6': str}.
# Dica 3: No dataset de população, converta o cod_munic e cod_uf para str no momento de carregar o dataset de população,
#         usando novamente o argumento dtype. Exatamente como foi feito no exercício da aula 3.
def merge_morbidade_pop(path_pop, path_morbidade):
    """retorne o merge dos datasets de população e morbidade.
    """
    # YOUR CODE HERE
    df_morbidade = pd.read_csv(path_morbidade, sep=";", 
                           header=0, decimal=",", 
                           na_values=["-"], 
                           low_memory=False, 
                           dtype={"cod_ibge6":"str"}
    )

    df_populacao = pd.read_excel(path_pop, header=0, dtype={"cod_uf":"str", "cod_munic":"str"})
    df_populacao["cod_ibge6"] = df_populacao["cod_uf"] + df_populacao["cod_munic"].str.slice(0,-1,1)
    return pd.merge(df_populacao, df_morbidade, on="cod_ibge6")

```

```python
## Faça seus testes aqui
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'

test_data = {'Internações': 149, 'AIH_aprovadas': 149, 'Valor_total': 101282.26}
df_morbidade_pop = merge_morbidade_pop(path_pop, path_morbidade)
df_00015 = df_morbidade_pop[(df_morbidade_pop['mes']=='Abr') & (df_morbidade_pop['cod_munic']=='00015')]
df_00015 = df_00015.loc[:,['Internações', 'AIH_aprovadas', 'Valor_total']]
dict2 = df_00015.to_dict(orient='records')[0]
dict1 = df_morbidade_pop.iloc[0, 8:11].to_dict()
dict1 == test_data or dict2 == test_data


 
```

Output:
```
True
```

<!-- grade_id: cell-0sb2b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
test_data = {'Internações': 149, 'AIH_aprovadas': 149, 'Valor_total': 101282.26}
df_morbidade_pop = merge_morbidade_pop(path_pop, path_morbidade)
df_00015 = df_morbidade_pop[(df_morbidade_pop['mes']=='Abr') & (df_morbidade_pop['cod_munic']=='00015')]
df_00015 = df_00015.loc[:,['Internações', 'AIH_aprovadas', 'Valor_total']]
dict2 = df_00015.to_dict(orient='records')[0]
dict1 = df_morbidade_pop.iloc[0, 8:11].to_dict()
assert dict1 == test_data or dict2 == test_data
```

```python

```