# Aula 3 - Exercícios - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
#Fazer o download dos datasets
# Se já tiver feito o download, não precisa fazer novamente.
# Se aparecer um erro dizendo que a pasta não está vazia é porque o download já foi feito
!git clone https://github.com/alexlopespereira/curso_ciencia_dados2020.git ../../curso_ciencia_dados2020
```

Output:
```
Cloning into '../../curso_ciencia_dados2020'...
remote: Enumerating objects: 323, done.[K
remote: Counting objects: 100% (323/323), done.[K
remote: Compressing objects: 100% (202/202), done.[K
remote: Total 323 (delta 125), reused 258 (delta 79), pack-reused 0
Receiving objects: 100% (323/323), 29.87 MiB | 10.08 MiB/s, done.
Resolving deltas: 100% (125/125), done.
Updating files: 100% (54/54), done.
```

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np
from numpy import nan as NA
import sqlalchemy as sqla

```

### 3.1 Carregue o arquivo do IDEB 2017 no formato de um DataFrame.  Crie uma função para calcular a média do IDEB de todos os municípios de um determinado Estado. 
#### Remova os registros cujo valor na coluna REDE sejam Municipal, Estadual ou Federal. Deixe somente os registros cujos valores na coluna REDE sejam Pública.


<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica 1: Antes de calcular, substitua valores '-' por np.nan (NA) com o metodo replace.
# Dica 2: Use a função loc para fazer o filtro de registros.
# Dica 3: Dentre várias maneiras possíveis, uma maneira bastante concisa é usar o operador & (and) para
# congregar dois critérios de filtro (seleção) de registros.
def media_ideb(df, sigla_estado):
    """retorne o valor da media do IDEB contido na coluna col do Estado especificado por sigla_estado.
    """
    # YOUR CODE HERE
    df = df.replace("-", NA)
    rede_filter = (df["REDE"] != 'Municipal') & (df["REDE"] != 'Estadual') & (df["REDE"] != 'Federal')

    df_filtered_series = df.loc[(df["SG_UF"] == sigla_estado) & rede_filter , "IDEB12_17"]

    return round(df_filtered_series.mean(),2)

```

```python
## Faça seus testes aqui
path_ideb='../../curso_ciencia_dados2020/data/originais/ideb/ideb_municipios2017.xlsx'
df_ideb = pd.read_excel(path_ideb, skiprows=8,skipfooter=3, header=1, na_values="-", dtype={'IDEB12_17': float})
#df_ideb
#res =  df_ideb["IDEB12_17"]
media_ideb(df_ideb, 'GO')

```

Output:
```
4.24
```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
###################################
### ATENÇÃO PARA ESTA INSTRUÇÃO ###
# Crie na célula anterior um dataframe df_ideb. 
# Use os argumentos skiprows e skipfooter para carregar a tabela corretamente.
###################################
assert media_ideb(df_ideb, 'DF') == 3.4
assert round(media_ideb(df_ideb, 'SP'),2) == 4.04
assert round(media_ideb(df_ideb, 'GO'),2) == 4.24
```

```python

```